Efficacy and mechanism of action of cipargamin as an antibabesial drug candidate

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #2 (Public review):

Summary:

In this manuscript, authors have tried to repurpose cipargamin (CIP), a known drug against Plasmodium and Toxoplasma against Babesia. They proved the efficacy of CIP on Babesia in nanomolar range. In silico analyses revealed the drug resistance mechanism through a single amino acid mutation at amino acid position 921 on the ATP4 gene of Babesia. Overall, the conclusions drawn by the authors are well justified by their data. I believe this study opens up a novel therapeutic strategy against babesiosis.

Strengths:

Authors have carried out a comprehensive study. All the experiments performed were carried out methodically and logically.

We appreciate your positive feedback. Your acknowledgment reinforces our commitment to rigor and thoroughness in our research.

Reviewer #3 (Public review):

Summary:

The authors aim to establish that cipargamin can be used for the treatment of infection caused by Babesia organisms.

Strengths:

The study provides strong evidence that cipargamin is effective against various Babesia species. In vitro growth assays were used to establish that cipargamin is effective against Babesia bovis and Babesia gibsoni. Infection of mice with Babesia microti demonstrated that cipargamin is as effective as the combination of atovaquone plus azithromycin. Cipargamin protected mice from lethal infection with Babesia rodhaini. Mutations that confer resistance to cipargamin were identified in the gene encoding ATP4, a P-type Na+ ATPase that is found in other apicomplexan parasites, thereby validating ATP4 as the target of cipargamin. A 7-day treatment of cipagarmin, when combined with a single dose of tafenoquine, was sufficient to eradicate Babesia microti in a mouse model of severe babesiosis caused by lack of adaptive immunity.

Thank you for the comments and for your time to review our manuscript.

Weaknesses:

Cipargamin was tested in vivo at a single dose administered daily for 7 days. Despite the prospect of using cipargamin for the treatment of human babesiosis, there was no attempt to identify the lowest dose of cipagarmin that protects mice from Babesia microti infection. In the SCID mouse model, cipargamin was tested in combination with tafenoquine but not with atovaquone and/or azithromycin, although the latter combination is often used as first-line therapy for human babesiosis caused by Babesia microti.

Thank you for your insightful comments. We agree that using a single daily dose over 7 days is one of the limitations in the in vivo trial. Our main goals were to demonstrate cipargamin's efficacy and understand its antibabesial agent mechanism. For future work, we plan to conduct dose‐optimization studies to determine the lowest effective dose in vivo. Regarding the drug combination in the SCID mouse model, although atovaquone and/or azithromycin are frequently used as first-line therapies for human babesiosis, resistance to these traditional drugs is emerging. Based on this challenge, we opted to evaluate a combination with tafenoquine as a novel partner, aiming to overcome resistance issues and improve therapeutic outcomes.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

None other than some minor grammatical mistakes.

We have corrected the grammatical mistakes.

Reviewer #3 (Recommendations for the authors):

The revised manuscript is much improved. I have the following comments.

Comment 1: Atovaquone plus azithromycin is effective against Babesia microti (Figure 1C) but not against Babesia rodhaini (Figure 1E). It would be valuable to provide a possible explanation.

Thank you for highlighting this issue. One potential explanation is that B. microti and B. rodhaini might have intrinsic differences in drug sensitivity and susceptibility. A previous study reported that both species possess a unique linear monomeric mitochondrial genome with a dual flip-flop inversion system, which generates four distinct genome structures (Hikosaka et al., 2012). In addition, previous studies have shown that mitochondria-associated energy production is greater in B. microti than in B. rodhaini (Shikano et al., 1998). This suggests that B. microti, whose metabolism is largely driven by mitochondrial function, may be more susceptible to drugs (like atovaquone) that induce parasite death by disrupting mitochondrial targets such as cytochrome b (Wormser et al., 2010). Moreover, B. rodhaini tends to proliferate more rapidly and causes acute infections, which may outpace any drug effects. Further, the rapid proliferation of apicomplexan parasites, as is the case in Plasmodium (Salcedo-Sora et al., 2014), Theileria (Metheni et al., 2015), and B. rodhaini (Rickard, 1970; Shikano et al., 1995), has been ascribed to glycolysis as the primary energy source. This may have contributed to the reduced efficacy of atovaquone and azithromycin in B. rodhaini-infected mice in the current study. Nonetheless, we plan to explore these interspecies differences in our future work.

Comment 2: The relapse that follows a 7-day treatment with cipargamin is transient in BALB/ mice infected with Babesia rodhaini (Figure 1E) but persistent in SCID mice infected with Babesia microti (Figure 5C). It would be valuable to provide a possible explanation.

Thank you for your insightful comment. One possible explanation is the difference in immune status between the two mouse models. BALB/c mice have a fully functional immune system that can likely clear residual parasites following a transient relapse after cipargamin treatment. In contrast, SCID mice lack an adaptive immune response, which might allow residual B. microti parasites to persist and cause a sustained relapse. Additionally, intrinsic differences between B. rodhaini and B. microti, such as growth rate or drug susceptibility, could also play a role. We plan to explore these factors in future studies.

Comment 3: The effect of cipargamin on parasite pH is the greatest when assessed 4 to 8 min after exposure is initiated (Figure 3E). Yet, resistance of parasites that carry a mutation in ATP4, the target of cipargamin, was assessed 20 min after cipargamin addition. At this time point, cipargamin has very little effect (Figure 3E). Accordingly, data reported in Figure 3G are of limited value.

Thank you for your comment. The initial pH increase we see around 4 to 8 minutes likely reflects the rapid inhibition of ATP4-mediated Na⁺/H⁺ exchange by cipargamin, which quickly alkalinizes the cell. However, after the initial increase, compensatory processes, such as proton influx or metabolic acid production, gradually restored the pH, resulting in a later decline. Although assessing the pH level at 20 minutes may have recorded less dramatic changes, it still allowed us to compare the sustained differences between wild-type and mutant strains. We agree that including earlier time points for the mutants might provide further insight and we will consider this in our future work.

Comment 4: In Figure 3H, please report the lack of statistical significance between wild-type parasites and parasites that carry the mutation L921V.

In Figure 3H, the ATPase activity in erythrocytes infected with wild-type parasites (6.31 ± 1.20 nmol Pi/mg protein/min) is higher than that of the L921V mutation (5.11 ± 0.50 nmol Pi/mg protein/min), but the difference is not statistically significant (P = 0.095), so no asterisk was added.

Comment 5: Tafenoquine was administered as a single 20 mg/kg dose. Please specify whether this dose is for tafenoquine succinate or tafenoquine base.

Thank you for raising this point. In our study, the single 20 mg/kg dose refers to tafenoquine succinate. We have clarified this detail in the revised manuscript (Line 40).

Comment 6: A single dose of 20 mg/kg tafenoquine succinate was first tested in the SCID mouse model of severe babesiosis by Mordue et al (JID 2019), not by Liu et al. (JID 2024). Please amend discussion accordingly (line 311). As correctly stated in the discussion, the single 20 mg/kg dose was not sufficient to prevent relapse of Babesia microti in the study by Mordue et al. Please provide a possible explanation for why no parasitemia was detected for 90 days in your SCID model (Figure 5C).

Thank you for your comment. We have modified the suggested citation (Line 309). As noted by Mordue et al. (JID 2019), a single 20 mg/kg dose of tafenoquine succinate was insufficient to prevent relapse in their SCID mouse model using B. microti (ATCC 30221 Gray strain). In our study, however, no parasitemia was detected for 90 days (Figure 5C) using the B. microti Peabody mjr strain (ATCC PRA-99). Differences in the parasite strain and the timing of treatment relative to infection may have contributed to the extended suppression of parasitemia observed in our study. We plan to explore these aspects in future work.

Comment 7: Real-time PCR was used to confirm eradication of Babesia microti infection (Figure 5D). Please specify the blood volume from which genomic DNA was extracted for each mouse. Please specify the amount of genomic DNA (i.e., not the volume) used in each reaction. Please explain how/why the cut-off was set at 35 cycles. What were the Ct values when blood was obtained from uninfected mice? For infected mice treated with cipargamin plus tafenoquine, there was no amplification. Was each reaction subjected to a maximum of 40 cycles (as suggested by Figure 5D)?

In our qPCR assay, genomic DNA was extracted from 200 µL of blood per mouse (Line 458). In each reaction, we used 100 ng of genomic DNA (Line 464), and the thermocycling conditions were set at 40 cycles. We set the cut-off at 35 cycles based on our optimization experiments: samples with Ct values ≤ 35 consistently indicated the presence of parasite DNA, while samples without parasite DNA (distilled water and DNA from uninfected mice) had CT values > 35 cycles or undetectable. Although each reaction was run for 40 cycles, for our analysis, we defined samples as negative if no signal was observed beyond cycle 35. In mice treated with cipargamin plus tafenoquine, no signal was detected until 40 cycles, indicating the absence of parasite DNA in the samples.

Comment 8: Persistence of parasite DNA in blood of tafenoquine treated mice highlights the limitation of PCR to assess persistence of infection. That is, PCR cannot distinguish between viable parasites and non-viable (or dead) parasites. An adoptive transfer of blood to immunocompromised mice can help determine whether persistence of DNA is due to persistence of viable parasites. Because the experiment was carried out in SCID mice, no adoptive transfer is needed. Few parasites are required for a successful infection of immunocompromised mice (SCID mice included). Given that parasitemia never rose following treatment of SCID mice with a single dose of tafenoquine, it is highly likely that parasite DNA detected on day 90 post-infection in these tafenoquine treated mice came from persistent non-viable/dead parasites.

We appreciate your comment and acknowledge that the use of PCR has limitations in differentiating between live and dead parasites. It is possible that the residual DNA may represent a small population of dormant parasites that are not actively replicating and thus remain below the detection threshold of parasitemia. Even in highly immunocompromised SCID mice, such dormant parasites might persist without causing overt infection under our experimental conditions. An adoptive transfer experiment in SCID mice, although not strictly necessary, could validate whether the detection of low levels of DNA comes from viable parasites capable of reactivating under different circumstances. Future studies using more sensitive viability assays or adoptive transfer approaches could provide further insights into this possibility.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Babesiosis is a disease brought on by intraerythrocytic parasites of the genus Babesia. Current chemotherapies are accompanied by side effects and parasite relapse. Therefore, it is crucial to develop highly effective drugs against Babesia. Cipargamin (CIP) has shown inhibition against apicomplexan parasites, mainly Plasmodium and Toxoplasma. This study evaluated the growth-inhibiting properties of CIP against Babesia spp. and investigated the mechanism of CIP on B. gibsoni. The half inhibitory concentration (IC50) values of CIP against the in vitro growth of B. bovis and B. gibsoni were 20.2 ± 1.4 nM and 69.4 ± 2.2 nM, respectively. CIP significantly inhibited the growth of B. microti and B. rodhaini in vivo. Resistance was conferred by L921V and L921I mutations in BgATP4, which reduced the sensitivity to CIP by 6.1- and 12.8-fold. The inhibitory potency of CIP against BgATP4-associated ATPase activity was moderately reduced in mutant strains, with a 1.3-fold and 2.4-fold decrease in BgATP4L921V and BgATP4L921I compared to that of BgATP4WT, respectively. An in silico investigation revealed reductions in affinity for CIP binding to BgATP4L921V and BgATP4L921I compared to BgATP4WT. Resistant strains showed no significant cross-resistance to atovaquone (ATO) or tafenoquine (TQ), with less than a onefold change in IC50 values. Combining CIP with TQ effectively eliminated B. microti infection in SCID mice, no relapse, and parasite DNA was not detected by qPCR within 90 days post-infection. Our findings reveal the efficacy of CIP as an anti-babesial agent, its limitations as a monotherapy due to resistance development, and the potential of combination therapy with TQ to overcome resistance and achieve complete parasite clearance.

Revised version: Babesiosis is a disease brought on by intraerythrocytic parasites of the genus Babesia. Current chemotherapies are accompanied by side effects and parasite relapse. Therefore, it is crucial to develop highly effective drugs against Babesia. Cipargamin (CIP) has shown inhibition against apicomplexan parasites, mainly Plasmodium and Toxoplasma. This study evaluated the growth-inhibiting properties of CIP against Babesia spp. and investigated the mechanism of CIP on B. gibsoni. The half inhibitory concentration (IC50) values of CIP against the in vitro growth of B. bovis and B. gibsoni were 20.2 ± 1.4 nM and 69.4 ± 2.2 nM, respectively. CIP significantly inhibited the growth of B. microti and B. rodhaini in vivo. Resistance was conferred by L921V and L921I mutations in BgATP4, which reduced the sensitivity to CIP by 6.1- and 12.8-fold. The inhibitory potency of CIP against BgATP4-associated ATPase activity was moderately reduced in mutant strains, with a 1.3-fold and 2.4-fold decrease in BgATP4L921V and BgATP4L921I, respectively, compared to that of BgATP4WT. An in silico investigation revealed reductions in affinity for CIP binding to BgATP4L921V and BgATP4L921I compared to BgATP4WT. Resistant strains showed no significant cross-resistance to atovaquone (ATO) or tafenoquine succinate (TQ), with less than a onefold change in IC50 values. Combining CIP with TQ effectively eliminated B. microti infection in SCID mice with no relapse, and parasite DNA was not detected by qPCR within 90 days post-infection. Our findings reveal the efficacy of CIP as an anti-babesial agent, its limitations as a monotherapy due to resistance development, and the potential of combination therapy with TQ to overcome said resistance and achieve complete parasite clearance.

Introduction

Babesia is an apicomplexan tick-transmitted hemoparasite, which not only impacts the livestock economy but also causes an emerging disease in humans (Jalovecka et al., 2019). Babesia is a global pathogen that is more prevalent in certain regions such as Asia, Europe, and North America. However, as climate change brings with it higher temperatures and humidity, ticks and their reservoir hosts are anticipated to expand northward for survival and activity (Gray and Ogden, 2021). The disease is known as babesiosis, commonly characterized by fever and hemolytic anemia, but chronic infections can be asymptomatic (Almazán et al., 2022). The fatality rate ranges from 1% among all cases to 3% among hospitalized cases, and as high as 20% in immunocompromised patients (Krause, 2019).

Currently, the most common drugs for the treatment of human babesiosis include a combination of atovaquone (ATO) plus azithromycin (AZI) or clindamycin (CLN) plus quinine (QUI) (Krause et al., 2021). Still, these recommended treatments often result in various problems. For instance, multiple mutations in the B. microti cytochrome b, which is targeted by ATO, were identified in patients with relapsing babesiosis (Holbrook et al., 2023; Krause et al., 2024; Lemieux et al., 2016; Marcos et al., 2022; Rogers et al., 2023; Rosenblatt et al., 2021; Simon et al., 2017). The combination of CLN and QUI is the last resort for patients with severe symptoms (Kletsova et al., 2017). Despite its efficacy, this combination can elicit adverse drug reactions (Vannier and Krause, 2012). The 8-aminoquinoline analog tafenoquine (TQ) was found to be effective in curing immunocompromised patients experiencing relapsing babesiosis caused by B. microti (Rogers et al., 2023). However, the efficacy of TQ treatment may vary between individuals and cases and its contraindication in patients with glucose-6-phosphate dehydrogenase (G6PD) deficiency limits its clinical use (Chu and Freedman, 2019). Due to the aforementioned issues, it is undoubtedly urgent to search for compounds that treat infections caused by Babesia spp. while simultaneously minimizing the detrimental side effects of antibabesial drugs.

Earlier version: Spiroindolone cipargamin (CIP), a promising antimalarial, has been found to effectively suppress the growth of all strains of Plasmodium falciparum and P. vivax with potency in the low nanomolar ranges (Rosling et al., 2018; Rottmann et al., 2010). The assessment of the drug indicated that CIP taken orally had good absorption, a long half-life, and exceptional bioavailability (Schmitt et al., 2022). Following oral CIP treatment (30 mg daily for three days) in adults with simple P. falciparum or P. vivax malaria, parasitemia was rapidly cleared in a phase II trial (Schmitt et al., 2022). Currently, a clinical trial is being conducted for the intravenous administration of CIP as a treatment for severe malaria patients (ClinicalTrials, 2024). CIP was also evaluated for treating toxoplasmosis as a second-line medicine in cases where there are intolerable toxicity issues or allergies to the currently used treatments. Mice infected with T. gondii that were treated with CIP on the day of infection and the following day had 90% fewer parasites five days post-infection (Zhou et al., 2014).

Revised version: Spiroindolone cipargamin (CIP), a promising antimalarial, has been found to effectively suppress the growth of all strains of Plasmodium falciparum and P. vivax with potency in the low nanomolar ranges (Rosling et al., 2018; Rottmann et al., 2010). The assessment of the drug indicated that CIP taken orally had good absorption, a long half-life, and exceptional bioavailability (Schmitt et al., 2022). Following oral CIP treatment (30 mg daily for three days) in adults with simple P. falciparum or P. vivax malaria, parasitemia was rapidly cleared in a phase II trial (Schmitt et al., 2022). Currently, a clinical trial is being conducted for the intravenous administration of CIP as a treatment for severe malaria patients (ClinicalTrials, 2024). CIP was also evaluated for treating toxoplasmosis as a second-line medicine in cases where there are intolerable toxicity issues or allergies to the currently used treatments. Mice infected with Toxoplasma gondii that were treated with CIP on the day of infection and the following day had 90% fewer parasites five days post-infection (Zhou et al., 2014).

PfATP4, a P-type ATPase in P. falciparum, functions as a Na⁺/H⁺ transporter critical for maintaining ionic balance (Spillman et al., 2013). Mutations in PfATP4 confer resistance to CIP, which disrupts Na⁺ homeostasis, leading to increased cytosolic Na⁺ concentration and various physiological changes, such as cytosolic alkalinization, osmotic swelling, etc. (Mohring et al., 2022). These mutations reduce the Na⁺-dysregulating effects of CIP and the resting Na⁺ levels. While PfATP4 is essential for P. falciparum survival, its homolog in T. gondii (TgATP4) is less critical, as T. gondii experiences brief, high Na⁺ exposure and can survive without TgATP4 expression (Lehane et al., 2019).

Due to its excellent efficacy against other apicomplexan parasites, including Plasmodium spp., and its ability to target parasite ATP4—a conserved protein essential for ion homeostasis in apicomplexan organisms—we hypothesize that CIP could also effectively inhibit Babesia spp., a related apicomplexan parasite. Therefore, the objectives of this study were to determine whether CIP could inhibit the growth of Babesia spp., namely B. bovis and B. gibsoni in vitro and B. microti and B. rodhaini in vivo, and identify the inhibitory mechanisms of CIP on Babesia parasites.

Results

Inhibitory efficacy of CIP on B. bovis and B. gibsoni in vitro

In vitro efficacy of CIP against B. bovis and B. gibsoni showed a steep growth inhibition curve with half inhibitory concentration (IC50) values of 20.2 ± 1.4 nM (Figure 1A) and 69.4 ± 2.2 nM (Figure 1B), respectively. The 50% cytotoxic concentration (CC50) value of CIP on Madin-Darby canine kidney (MDCK) cells and human foreskin fibroblasts (HFF) were 38.7 ± 2.0 μM and 70.8 ± 4.9 μM (Figure supplement 1), respectively. Based on these values, the predicted selectivity indices (SIs), which reflect the drug’s safety and specificity, were calculated to be greater than 500. Furthermore, at a concentration of 100 μM, CIP exhibited a low erythrocyte hemolysis rate of 0.11 ± 0.03 % (data not shown).

CIP demonstrates potent inhibition on Babesia spp.

Earlier version: (A) and (B) Dose-dependent growth curve of CIP on B. bovis and B. gibsoni in vitro. Each value represents the mean ± standard deviation (SD) of three independent experiments carried out in triplicate. IC50: the half maximal inhibitory concentration. (C) Inhibitory effects of CIP and atovaquone (ATO) plus azithromycin (AZI) on the proliferation of B. microti in BALB/c mice. (D) Hematocrit (HCT) values in mice treated with CIP or ATO plus AZI compared with vehicle-treated mice. (E) Inhibitory effects of CIP and ATO plus AZI on the proliferation of B. rodhaini in BALB/c mice. (F) Survival rates of CIP-treated, ATO plus AZI-treated, and vehicle-treated mice. The treatment time is shown by two-way arrows, and significant differences (P < 0.01) between the drug-treated groups and the vehicle-treated control group are indicated by asterisks. The data from one of six individual experiments are expressed as means ± SD. *, P ˂ 0.05; **, P ˂ 0.01.

Revised version: (A) and (B) Dose-dependent growth curve of CIP on B. bovis and B. gibsoni in vitro. Each value represents the mean ± standard deviation (SD) of three independent experiments carried out in triplicate. IC50: the half maximal inhibitory concentration. (C) Inhibitory effects of CIP and atovaquone (ATO) plus azithromycin (AZI) on the proliferation of B. microti in BALB/c mice. (D) Hematocrit (HCT) values in mice treated with CIP or ATO plus AZI compared with vehicle-treated mice. (E) Inhibitory effects of CIP and ATO plus AZI on the proliferation of B. rodhaini in BALB/c mice. (F) Survival rates of CIP-treated, ATO plus AZI-treated, and vehicle-treated mice. The treatment time is shown by two-way arrows, and significant differences (P < 0.01) between the drug-treated groups and the vehicle-treated control group are indicated by asterisks. The data from one of six individual experiments are expressed as means ± SD. **, P ˂ 0.01.

CIP effect on B. microti and B. rodhaini infections in vivo

Earlier version: Concurrently, CIP showed effective inhibition on B. microti and B. rodhaini in vivo. The parasitemia of B. microti-infected BALB/c mice increased dramatically in the vehicle-treated control group and peaked at 10 days postinfection (DPI) (38.55 ± 4.32%) (Figure 1C). On the other hand, seven days of treatment with CIP (20 mg/kg) or atovaquone (ATO) plus azithromycin (AZI) administered orally resulted in a significantly lower peak parasitemia, 1.06 ± 0.20% and 1.61 ± 0.20%, respectively (Figure 1C). Hematocrit (HCT) variations were tracked every 4 days as an indicator of anemia in B. microti-infected mice. The vehicle-treated group showed a drop in HCT levels at 12 DPI and 16 DPI (P < 0.01) (Figure 1D). No significant reduction in HCT levels was observed in the CIP-treated group or the ATO plus AZI-treated group (Figure 1D). This indicates that the administration of CIP could control B. microti infection and prevent anemia from developing in B. microti-infected mice. BALB/c mice infected with B. rodhaini treated with sesame oil or ATO plus AZI showed high parasitemia, 90.73 ± 1.97%, and 86.23 ± 3.06%, respectively (Figure 1E), and all mice died within 8 DPI (Figure 1F). CIP treatment in B. rodhaini-infected mice precluded the emergence of parasitemia for the following eight days (Figure 1E), which led to 66.67% of mice surviving the challenge infection (Figure 1F). At 12 DPI, parasites had recurred in all CIP-treated B. rodhaini-infected mice (10.32 ± 15.51%), which were eventually cleared as indicated by undetectable parasites at 18 DPI (Figure 1E).

Revised version: Concurrently, CIP showed effective inhibition on B. microti and B. rodhaini in vivo. The parasitemia of B. microti-infected BALB/c mice increased dramatically in the vehicle-treated control group and peaked at 10 days postinfection (DPI) (38.55 ± 4.32%) (Figure 1C). On the other hand, seven days of treatment with CIP (20 mg/kg) or atovaquone (ATO) plus azithromycin (AZI) administered orally resulted in a significantly lower peak parasitemia, 1.06 ± 0.20% and 1.61 ± 0.20%, respectively (Figure 1C). Hematocrit (HCT) variations were tracked every 4 days as an indicator of anemia in B. microti-infected mice. The vehicle-treated group showed a drop in HCT levels at 12 DPI and 16 DPI (P < 0.01) (Figure 1D). No significant reduction in HCT levels was observed in the CIP-treated group or the ATO plus AZI-treated group (Figure 1D). This indicates that the administration of CIP could control B. microti infection and prevent anemia from developing in B. microti-infected mice. BALB/c mice infected with B. rodhaini treated with sesame oil or ATO plus AZI showed high parasitemia, 90.73 ± 1.97%, and 86.23 ± 3.06%, respectively (Figure 1E), and all mice died within 7 DPI (Figure 1F). CIP treatment in B. rodhaini-infected mice precluded the emergence of parasitemia for the following eight days (Figure 1E), which led to 66.67% of mice surviving the challenge infection (Figure 1F). At 12 DPI, parasites had recurred in all CIP-treated B. rodhaini-infected mice (10.32 ± 15.51%), which were eventually cleared as indicated by undetectable parasites at 18 DPI (Figure 1E).

Identification of B. gibsoni ATP4 mutation in CIP-resistant strains

Earlier version: After being exposed to CIP at increasing concentrations up to 10 times the IC50, the resistant parasites in two of the culture wells were able to regrow. We sequenced the B. gibsoni ATP4 gene from the wild-type and two resistant strains. The wild-type strain has a C at nucleotide 2,761, which translates to leucine (Figure 2A). In one resistant strain, a single nucleotide variant (SNV) in BgATP4 with a substitution at position 2,761 (from C to G) was found − a nonsynonymous coding change from leucine to valine (L921V) (Figure 2B). In another resistant strain, the mutation occurred in the same position. However, the nucleotide substitution was from C to A, and the coding changed from leucine to isoleucine (L921I) (Figure 2C). Next-generation sequencing (NGS) revealed that for BgATP4L921V, 99.97% of 7,960 reads were G at nucleotide 2,761, and for BgATP4L921I, 99.92% of 7,862 reads were A at nucleotide 2,761 (Figure 2D). BgATP4L921V and BgATP4L921I lines were tested for their susceptibility to CIP, and had IC50 values of 421.0 ± 15.9 nM and 887.9 ± 62.0 nM, respectively (Figure 2E). These findings demonstrate a 6.1- and 12.8-fold reduction in CIP sensitivity of the resistant parasite lines BgATP4L921V and BgATP4L921I.

Revised version: After being exposed to CIP at increasing concentrations up to 10 times the IC50, the resistant parasites in two of the culture wells were able to regrow. We sequenced the B. gibsoni ATP4 gene from the wild-type and two resistant strains. The wild-type strain has a C at nucleotide 2,761, which translates to leucine (Figure 2A). In one resistant strain, a single nucleotide variant (SNV) in BgATP4 with a substitution at position 2,761 (from C to G) was found − a nonsynonymous coding change from leucine to valine (L921V) (Figure 2B). In another resistant strain, the mutation occurred in the same position. However, the nucleotide substitution was from C to A, and the coding changed from leucine to isoleucine (L921I) (Figure 2C). Next-generation sequencing (NGS) revealed that for Bgatp42761C>G, 99.97% of 7,960 reads were G at nucleotide 2,761, and for Bgatp42761C>A, 99.92% of 7,862 reads were A at nucleotide 2,761 (Figure 2D). BgATP4L921V and BgATP4L921I lines were tested for their susceptibility to CIP, and had IC50 values of 421.0 ± 15.9 nM and 887.9 ± 62.0 nM, respectively (Figure 2E). These findings demonstrate a 6.1- and 12.8-fold reduction in CIP sensitivity of the resistant parasite lines BgATP4L921V and BgATP4L921I.

Mutations in BgATP4 mediate CIP resistance.

(A-C) Representative sequencing chromatogram of wild-type and resistant parasites from CIP-treated B. gibsoni. The resistant parasite genomic DNA is extracted from blood samples after a 60 day-treatment. The BgATP4 gene was amplified and sequenced using the DNA. (D) Genes of high-frequency sequence variants detected by NGS. (E) Dose-dependent growth curve of BgATP4WT, BgATP4L921V, and BgATP4L921I in vitro. Each value represents the mean ± standard deviation (SD) of three independent experiments carried out in triplicate.

The effect of CIP on BgATP4WT, BgATP4L921V and BgATP4L921I function

Earlier version: Microscopic observation of thin blood smears was performed to determine the morphological changes of B. gibsoni exposed to CIP. The CIP-treated parasites became swollen after incubation with the drug for 72 h (Figure 3A). For both the treatment and control groups, one hundred parasites were measured. The mean size of treated parasites was notably bigger than the parasites in the untreated group (P ˂ 0.0001) (Figure 3B). Significant vacuolization was observed in the cytoplasm of parasites in the CIP-treated group, as revealed by transmission electron microscopy (TEM). Despite this, the nuclear membrane structure and parasitic membranes remained intact until the parasites were completely destroyed (Figure 3C). The addition of the ATP4 inhibitor CIP resulted in a time-dependent increase in the concentrations of [Na+]i in wild-type B. gibsoni, with improved signal-to-noise ratios at the higher drug concentration of 20 nM (Figure 3D). We observed that the Na+ concentrations in both BgATP4L921V and BgATP4L921I lines were lower when compared with those of the control BgATP4WT line after being exposed to 20 nM CIP for 20 min, with a significantly lower Na+ concentration in BgATP4L921I (P = 0.0087) (Figure 3F). We also demonstrated here that the addition of CIP in wild-type B. gibsoni caused an increase in the cytosolic pH (Figure 3E). Specifically, 4 min after the drug was added, the average pH of 20 nM CIP group reached as high as 7.278, while the 1 nM CIP group reached 7.089 and the untreated group reached 7.062 (Figure 3E). The pH values of the 20 nM CIP group were consistently higher than those of the other two groups, although declining with time (Figure 3E). In resistant lines, a 20-min exposure to 20 nM CIP caused small changes in the pH values compared to the wild-type line, with the BgATP4L921I line (7.048 ± 0.042) having a notably lower pH value (P = 0.0229) (Figure 3G).

Revised version: Microscopic observation of thin blood smears was performed to determine the morphological changes of B. gibsoni exposed to CIP. The CIP-treated parasites became swollen after incubation with the drug for 72 h (Figure 3A). For both the treatment and control groups, one hundred parasites were measured. The mean size of treated parasites was notably bigger than the parasites in the untreated group (P ˂ 0.0001) (Figure 3B). Significant vacuolization was observed in the cytoplasm of parasites in the CIP-treated group, as revealed by transmission electron microscopy (TEM). Despite this, the nuclear membrane structure and parasitic membranes remained intact until the parasites were completely destroyed (Figure 3C). The addition of the ATP4 inhibitor CIP resulted in a time-dependent increase in the concentrations of [Na+]i in wild-type B. gibsoni, with improved signal-to-noise ratios at the higher drug concentration of 20 nM (Figure 3D). We observed that the Na+ concentrations in both BgATP4L921V and BgATP4L921I lines were lower when compared with those of the control BgATP4WT line after being exposed to 20 nM CIP for 20 min, with a significantly lower Na+ concentration in BgATP4L921I (P = 0.0087) (Figure 3F). We also demonstrated here that the addition of CIP in wild-type B. gibsoni caused an increase in the cytosolic pH (Figure 3E). Specifically, 4 min after the drug was added, the average pH of the 20 nM CIP group reached as high as 7.278, while the 1 nM CIP group reached 7.089 and the untreated group reached 7.062 (Figure 3E). The pH values of the 20 nM CIP group were consistently higher than those of the other two groups, although declining with time (Figure 3E). In resistant lines, a 20-min exposure to 20 nM CIP caused small changes in the pH values compared to the wild-type line, with the BgATP4L921I line (7.048 ± 0.042) having a notably lower pH value (P = 0.0229) (Figure 3G).

Mechanistic basis for resistance to CIP conferred by the L921V and L921I mutations in BgATP4.

Earlier version: (A) Untreated and CIP-treated parasite morphology after incubation for 72 h. Scale bar: 5 µm. (B) Sizes of 100 parasites in two groups measured with ImageJ software in panels A. Statistically significant differences between the means of variables determined by t-test. ****, P ˂ 0.0001. (C) TEM of untreated and CIP-treated parasite. N, Nucleus; SB, spherical body; V, vacuole. Scale bar: 500 nm. (D) [Na+]i concentrations after addition of CIP in BgATP4WT line. Representative traces from the experiment that highlight the impact of adding 20 nM CIP (blue), 1 nM CIP (green), or 0 nM CIP (grey) on the concentration [Na+]i of the BgATP4WT line. (E) Alkalinization of pHi in BgATP4WT line upon addition of the ATP4 inhibitor. (F) Addition of 20 nM CIP to the wild-type and resistant parasite lines results in different [Na+]i concentrations. (G) Addition of 20 nM CIP to the wild-type and resistant parasite lines results in different pHi concentrations. Experiments were performed in technical duplicates for at least three biological repeats. (H) Data acquired in the low Na+ condition (containing only the 2 mM Na+ introduced upon addition of 1 mM Na2ATP) was subtracted from data obtained in the high Na+ condition to determine the ATPase activity related to the BgATP4 proteins. The results are presented as the average of data from three independent tests. (I) Dose-dependent BgATP4-associated ATPase activity curve of BgATP4WT, BgATP4L921V, and BgATP4L921I in vitro. ATPase activity was determined at pH 7.2 in the presence of 150 mM Na+ and 1mM Na2ATP. Each value represents the mean ± standard deviation (SD) of three independent experiments carried out in triplicate. *, P ˂ 0.05; **, P ˂ 0.01.

Revised version: (A) Untreated and CIP-treated parasite morphology after incubation for 72 h. Scale bar: 5 µm. (B) Sizes of 100 parasites in two groups measured with ImageJ software in panels A. Statistically significant differences between the means of variables determined by t-test. ****, P ˂ 0.0001. (C) TEM of untreated and CIP-treated parasite. N, Nucleus; SB, spherical body; V, vacuole. Scale bar: 500 nm. (D) [Na+]i concentrations after the addition of CIP in the BgATP4WT line. Representative traces from the experiment that highlight the impact of adding 20 nM CIP (blue), 1 nM CIP (green), or 0 nM CIP (grey) on the concentration [Na+]i of the BgATP4WT line. (E) Alkalinization of pHi in BgATP4WT line upon addition of the ATP4 inhibitor. (F) Addition of 20 nM CIP to the wild-type and resistant parasite lines results in different [Na+]i concentrations. (G) Addition of 20 nM CIP to the wild-type and resistant parasite lines results in different pHi concentrations. Experiments were performed in technical duplicates for at least three biological repeats. (H) Data acquired in the low Na+ condition (containing only the 2 mM Na+ introduced upon the addition of 1 mM Na2ATP) was subtracted from data obtained in the high Na+ condition to determine the ATPase activity related to the BgATP4 proteins. The results are presented as the average of data from three independent tests. (I) Dose-dependent BgATP4-associated ATPase activity curve of BgATP4WT, BgATP4L921V, and BgATP4L921I in vitro. ATPase activity was determined at pH 7.2 in the presence of 150 mM Na+ and 1 mM Na2ATP. Each value represents the mean ± standard deviation (SD) of three independent experiments carried out in triplicate. *, P ˂ 0.05; **, P ˂ 0.01.

Sensitivity of BgATP4-associated ATPase activity to CIP in BgATP4WT, BgATP4L921V, and BgATP4L921I

Earlier version: The BgATP4-associated ATPase activity in erythrocytes infected with BgATP4WT (6.31 ± 1.20 nmol Pi/mg protein/min), measured in the presence of 150 mM Na+, was higher than that observed in BgATP4L921V (5.11 ± 0.50 nmol Pi/mg protein/min) and BgATP4L921I (4.58 ± 0.53 nmol Pi/mg protein/min) (P = 0.04) (Figure 3H).

Revised version: The BgATP4-associated ATPase activity in erythrocytes infected with BgATP4WT (6.31 ± 1.20 nmol Pi/mg protein/min), measured in the presence of 150 mM Na+, was higher than those observed in BgATP4L921V (5.11 ± 0.50 nmol Pi/mg protein/min) and BgATP4L921I (4.58 ± 0.53 nmol Pi/mg protein/min) (P = 0.04) (Figure 3H).

We further investigated the concentration-dependent inhibition of BgATP4-associated ATPase activity by CIP in wild-type and mutant parasites. In membranes prepared from B.gibsoni, CIP inhibited BgATP4-associated ATPase activity with IC50 values of 111.4 ± 31.8 nM for BgATP4WT, 149.8 ± 21.7 nM for BgATP4L921V, and 269.5 ± 29.8 nM for BgATP4L921I. The potency of CIP in inhibiting BgATP4-associated ATPase activity was reduced by 1.3-fold and 2.4-fold in membranes prepared from BgATP4L921V and BgATP4L921I, respectively, compared to BgATP4WT (Figure 3I).

Multiple sequence alignment of Babesia ATP4 and molecular docking

The whole amino acid sequence of B. gibsoni ATP4 (GenBank: KAK1443404.1) shared identity values of 29.75%, 49.40%, 49.67%, 62.21%, and 52.47% with Homo sapiens ATP4 (GenBank: NM_000704.3), P. falciparum ATP4 (GenBank: PF3D7_1211900), T. gondii ATP4 (GenBank: XP_018635122.1), B. bovis ATP4 (PiroplasmaDB: BBOV_IV010020), and B. microti ATP4 (GenBank: BMR1_03g01005), respectively (Figure supplement 2).

The pLDDT (predicted Local Distance Difference Test) value of BgATP4WT prediction was 80.7 using Colab-fold. Multiple potential binding sites for CIP were revealed by blind docking throughout the whole protein surface (Figure supplement 3). CIP binds in close proximity to L921, as demonstrated by focused docking on this area (Figure 4A). The contribution of each residue to the predicted binding affinity in either mutant structure was reduced; the precise values of BgATP4WT (Figure 4B), BgATP4L921V (Figure 4C), and BgATP4L921I (Figure 4D) were -6.43, -6.40 and -6.26 kcal/mol, respectively. The interactions of CIP from each docking simulation are shown in Table supplement 3.

Binding sites proximal to BgATP4 residue 921 predicted by molecular docking.

(A) The lowest energy poses for CIP were located in reference to the whole protein structure, docking against the WT (green), L921V (yellow), and L921I (pink) mutant BgATP4. The side chain of L921 is also shown in a red stick at its position. (B-D) The zoomed views of the binding locations of CIP.

Cross-resistance of BgATP4L921V and BgATP4L921I mutants to ATO and TQ

The BgATP4WT, BgATP4L921V, and BgATP4L921I lines were tested for their susceptibility to ATO and TQ. The IC50 values for these lines were 380.1 ± 3.1 nM, 412.9 ± 5.4 nM, and 360.3 ± 7.9 nM, respectively, for ATO (Figure 5A), and 39.5 ± 0.4 μM, 29.9 ± 1.4 μM, and 39.3 ± 0.3 μM, respectively, for TQ (Figure 5B). The IC50 values for both ATO and TQ in the resistant strains showed only slight changes compared to the wild-type strain, with less than a onefold difference. This minimal variation suggests that the resistant strain has a mild alteration in susceptibility to ATO and TQ, but not enough to indicate strong resistance or significant cross-resistance.

Cross-resistance between ATO and TQ in resistant parasites and combination therapy based on CIP plus TQ.

Earlier version: (A) Dose-dependent growth curve of BgATP4WT, BgATP4L921V, and BgATP4L921I by ATO treatment in vitro. (B) Dose-dependent growth curve of BgATP4WT, BgATP4L921V, and BgATP4L921I by TQ treatment in vitro. Each value represents the mean ± standard deviation (SD) of three independent experiments carried out in triplicate. (C) Inhibitory effects of CIP plus TQ on the proliferation of B. microti in SCID mice. The data are presented as the means from one of five independent experiments. (D) Parasite DNA was detected by qPCR on genomic DNA extracted from blood collected from untreated and treated SCID mice infected with B. microti at 90 DPI. A dotted grey line across the graph represents the average cut-off Cq value. Cut-off Cq ≤ 35 was considered as positive, while Cq > 35 or no amplification was considered as negative. Each sample was analyzed in duplicate, and two independent experiments were conducted. NA indicates no amplification.

Revised version: (A) Dose-dependent growth curve of BgATP4WT, BgATP4L921V, and BgATP4L921I by ATO treatment in vitro. (B) Dose-dependent growth curve of BgATP4WT, BgATP4L921V, and BgATP4L921I by TQ treatment in vitro. Each value represents the mean ± standard deviation (SD) of three independent experiments carried out in triplicate. (C) Inhibitory effects of CIP plus TQ on the proliferation of B. microti in SCID mice. Treatment started at 4 DPI: CIP was given at 20 mg/kg once daily for 7 days, TQ was administered as a single 20 mg/kg dose, and the combination group received both treatments (CIP at 20 mg/kg once daily for 7 days plus a single dose of TQ at 20 mg/kg). The data are presented as the means from one of five independent experiments. (D) Parasite DNA was detected by qPCR on genomic DNA extracted from blood collected from untreated and treated SCID mice infected with B. microti at 90 DPI. A dotted grey line across the graph represents the average cut-off Cq value. Cut-off Cq ≤ 35 was considered positive, while Cq > 35 or no amplification was considered negative. Each sample was analyzed in duplicate, and two independent experiments were conducted. NA indicates no amplification.

Combination treatment of CIP plus TQ in SCID mice with B. microti infection

Earlier version: The parasitemia of B. microti-infected SCID mice in the vehicle group increased dramatically, peaking at 10 DPI (77.03 ± 2.45%) (Figure 5C). Although the trend showed a subsequent decline, the parasitemia remained around 50% until the mice were euthanized at 90 DPI. Treatment with CIP (20 mg/kg) initially reduced parasitemia to undetectable by 18 DPI. However, a relapse occurred, with parasitemia increasing rapidly and stabilizing at levels comparable to the vehicle group during the subsequent observation period. Notably, no mutations were detected in the relapsed B. microti parasites from SCID mice (data not shown). In contrast, parasitemia was completely cleared in the TQ and CIP plus TQ groups by 8 DPI and 6 DPI, respectively, with no parasites observed under the microscope thereafter. Absolute qPCR analysis at 90 DPI detected parasite DNA in all groups except the CIP plus TQ group (Figure 5D).

Revised version: The parasitemia of B. microti-infected SCID mice in the vehicle group increased dramatically, peaking at 10 DPI (77.03 ± 2.45%) (Figure 5C). Although the trend showed a subsequent decline, the parasitemia remained around 50% until the mice were euthanized at 90 DPI. Treatment with CIP (20 mg/kg) initially reduced parasitemia to undetectable by 18 DPI. However, a relapse occurred, with parasitemia increasing rapidly and stabilizing at levels comparable to the vehicle group during the subsequent observation period. Notably, no mutations were detected in the relapsed B. microti parasites from SCID mice (data not shown). In contrast, parasitemia was completely cleared in the TQ and CIP plus TQ groups by 8 DPI and 6 DPI, respectively, with no parasites observed under the microscope thereafter. qPCR analysis at 90 DPI detected parasite DNA in all groups except the CIP plus TQ group (Figure 5D).

Discussion

Earlier version: The repositioning of antimalarial drugs is critical in developing novel strategies for treating babesiosis. CIP is a novel compound that inhibits Plasmodium development by targeting ATP4 and has been extensively tested in Phase 1 and Phase 2 clinical trials (Qiu et al., 2022). Since ATP4 is conserved across apicomplexans, including Babesia species, our study aimed to repurpose the antimalarial CIP by assessing its efficacy on Babesia species (Lehane et al., 2019). The IC50 values of CIP against B. bovis and B. gibsoni in vitro were lower than the previously reported IC50 value of TQ (Carvalho et al., 2020; Ji et al., 2022b) and ATO against B. gibsoni (Matsuu et al., 2004). The present investigation also demonstrated that the inhibitory effects of CIP on B. microti-infected BALB/c mice were comparable to that of ATO plus AZI, the combination recommended by the CDC in the United States. CIP was also proven to protect mice from the deadly B. rodhaini infection, with survival up to 67%, as recorded in the current trial. These results suggest that CIP may be a potential chemotherapy candidate for babesiosis.

Revised version: The repositioning of antimalarial drugs is critical in developing novel strategies for treating babesiosis. CIP is a novel compound that inhibits Plasmodium development by targeting ATP4 and has been extensively tested in Phase 1 and Phase 2 clinical trials (Qiu et al., 2022). Since ATP4 is conserved across apicomplexans, including Babesia species, our study aimed to repurpose the antimalarial CIP by assessing its efficacy on Babesia species (Lehane et al., 2019). The IC50 values of CIP against B. bovis and B. gibsoni in vitro were lower than the previously reported IC50 value of TQ (Carvalho et al., 2020; Ji et al., 2022b) and ATO against B. gibsoni (Matsuu et al., 2004). The present investigation also demonstrated that the inhibitory effects of CIP on B. microti-infected BALB/c mice were comparable to that of ATO plus AZI, the combination recommended by the CDC in the United States. CIP was also proven to protect mice from the deadly B. rodhaini infection, with a survival rate of up to 67%, as recorded in the current trial. These results suggest that CIP may be a potential chemotherapy candidate for babesiosis.

Earlier version: In a previous study, the P. falciparum Dd2 strain that acquired resistance to CIP carried the G358S mutation in the PfATP4 protein, a P. falciparum P-type Na+ ATPase (Qiu et al., 2022). The vital protein ATP4 is found in the parasite plasma membrane and is specific to the subclass of apicomplexan parasites (Mohring et al., 2022). To date, in vitro evolution experiments using CIP have produced at least 18 parasite lines with various PfATP4 mutations (Lee and Fidock, 2016; Rottmann et al., 2010). In another study involving T. gondii, a cell line that carried the mutation G419S in the TgATP4 gene was 34 times less susceptible to CIP than that of TgATP4WT (Qiu et al., 2022). It follows that there is a significant chance that the resistant Babesia parasites will also emerge from CIP exposure. In this study, we successfully produced two CIP-resistant strains using six independent selections. The newly discovered L921V and L921I mutations in BgATP4 decreased the CIP sensitivity by 6.1 and 12.8 times, respectively. We provided compelling evidence that natural variety mutations L921V and L921I in BgATP4 significantly affected the protein’s susceptibility to CIP inhibition, albeit the mutational sites were different from those of P. falciparum and T. gondii. Based on our observation, the growth and generation rates of the mutant strains are comparable to those of the wild-type strain.

Revised version: In a previous study, the P. falciparum Dd2 strain that acquired resistance to CIP carried the G358S mutation in the PfATP4 protein, a P. falciparum P-type Na+ ATPase (Qiu et al., 2022). The vital protein ATP4 is found in the parasite plasma membrane and is specific to the subclass of apicomplexan parasites (Mohring et al., 2022). To date, in vitro evolution experiments using CIP have produced at least 18 parasite lines with various PfATP4 mutations (Lee and Fidock, 2016; Rottmann et al., 2010). In another study involving T. gondii, a cell line that carried the mutation G419S in the TgATP4 gene was 34 times less susceptible to CIP than that of TgATP4WT (Qiu et al., 2022). It follows that there is a significant possibility that resistant Babesia parasites will emerge from CIP exposure. In this study, we successfully produced two CIP-resistant strains using six independent selections. The newly discovered L921V and L921I mutations in BgATP4 decreased the CIP sensitivity by 6.1 and 12.8 times, respectively. We provided compelling evidence that natural variety mutations L921V and L921I in BgATP4 significantly affected the protein’s susceptibility to CIP inhibition, albeit the mutational sites were different from those of P. falciparum and T. gondii. Based on our observation, the growth and generation rates of the mutant strains are comparable to those of the wild-type strain.

Earlier version: Although ATP4 was initially recognized as a Ca2+ transporter (Krishna et al., 2001), current evidence suggests that ATP4 functions as an ATPase for exporting Na+ while importing H+ (Mohring et al., 2022), as well as causing a variety of other physiological perturbations including an increase in the volume of parasites and infected erythrocytes, due to the osmotic impact of the [Na+]cyt increase (Dennis et al., 2018), a decline in cholesterol extrusion from the parasite plasma membrane as a result of the increase in [Na+]cyt (Das et al., 2016), and an intensified rigidity of erythrocytes infected with ring-stage parasites (Zhang et al., 2016). In this study, the CIP-exposed wild-type B. gibsoni became swollen, which was identical to a prior study on P. falciparum (Dennis et al., 2018). Interestingly, TEM analysis of parasites incubated with CIP revealed ultrastructural alterations characterized by significant vacuole formation in the cytoplasm, a hallmark of stress or early stages of cell death. These changes were similar to those observed in B. bovis treated with clotrimazole and ketoconazole, which ultimately led to growth inhibition or parasite death (Bork et al., 2003). Despite the vacuolization, the structural integrity of the nucleus and membranes was maintained, suggesting a gradual process where vacuolization precedes the complete destruction of the parasite. One explanation is the swelling of the isolated parasites, which can be ascribed to the osmotic consequences of Na+ uptake and is contingent upon the presence of Na+ in the external environment.

Revised version: Although ATP4 was initially recognized as a Ca2+ transporter (Krishna et al., 2001), current evidence suggests that ATP4 functions as an ATPase for exporting Na+ while importing H+ (Mohring et al., 2022), as well as causing a variety of other physiological perturbations including an increase in the volume of parasites and infected erythrocytes due to the osmotic impact of the [Na+]cyt increase (Dennis et al., 2018), a decline in cholesterol extrusion from the parasite plasma membrane as a result of the increase in [Na+]cyt (Das et al., 2016), and an intensified rigidity of erythrocytes infected with ring-stage parasites (Zhang et al., 2016). In this study, the CIP-exposed wild-type B. gibsoni became swollen, which was identical to a prior study on P. falciparum (Dennis et al., 2018). Interestingly, transmission electron microscopy (TEM) analysis of parasites incubated with CIP revealed ultrastructural alterations characterized by significant vacuole formation in the cytoplasm, a hallmark of stress or early stages of cell death. These changes were similar to those observed in B. bovis treated with clotrimazole and ketoconazole, which ultimately led to growth inhibition or parasite death (Bork et al., 2003). Despite the vacuolization, the structural integrity of the nucleus and membranes was maintained, suggesting a gradual process where vacuolization precedes complete destruction of the parasite. One explanation is the swelling of the isolated parasites, which can be ascribed to the osmotic consequences of Na+ uptake and is contingent upon the presence of Na+ in the external environment.

Earlier version: The output of Na+ and input of H+ diminished upon CIP-induced inhibition of PfATP4, and the continued outflow of H+ via V-type H+-ATPase led to an alkalinization that ultimately killed the parasites (Spillman et al., 2013). Our capacity to measure a time-dependent increase in the concentration of [Na+]i and pH value for wild-type B. gibsoni facilitated us to gain insight into the basic mechanisms of CIP on BgATP4WT function. Furthermore, our results validate that internal alkalinization was the main factor in Babesia death and support our hypothesis that the swollen isolated parasites were produced by Na+ absorption. To explore further how mutations in BgATP4 are associated with the upregulation of the parasite’s [Na+]i and [H+]i, we tested two BgATP4-mutant lines that were chosen previously with BgATP4 inhibitors (BgATP4L921V and BgATP4L921I). Due to the presence of L921V and L921I mutations in drug-resistant strains of BgATP4, the concentration of Na+ did not increase as much as it would have in the wild-type strain following the addition of 20 nM CIP for 20 min. As intraerythrocytic alkalinization rises, the same outcome happens. From these results, we deduced how natural mutations in BgATP4 may affect ATP4 inhibitor susceptibility by dysregulating H+ and Na+ balance, which helped parasites survive in a relatively high concentration of CIP. An illustration of the putative processes for regulating Na+ and H+ in erythrocytes infected with the Babesia parasite is presented in Figure supplement 4.

Revised version: The output of Na+ and input of H+ diminished upon CIP-induced inhibition of PfATP4, and the continued outflow of H+ via V-type H+-ATPase led to an alkalinization that ultimately killed the parasites (Spillman et al., 2013). Our capacity to measure a time-dependent increase in the concentration of [Na+]i and pH value for wild-type B. gibsoni facilitated us to gain insight into the basic mechanisms of CIP on BgATP4WT function. Furthermore, our results corroborate internal alkalinization as the main factor in Babesia death and support our hypothesis that the swollen isolated parasites were produced by Na+ absorption. To explore further how mutations in BgATP4 are associated with the upregulation of the parasite’s [Na+]i and [H+]i, we tested two BgATP4-mutant lines that were chosen previously with BgATP4 inhibitors (BgATP4L921V and BgATP4L921I). Due to the presence of L921V and L921I mutations in drug-resistant strains of BgATP4, the concentration of Na+ did not increase as much as it would have in the wild-type strain following the addition of 20 nM CIP for 20 min. As intraerythrocytic alkalinization rises, the same outcome happens. From these results, we deduced how natural mutations in BgATP4 may affect ATP4 inhibitor susceptibility by dysregulating H+ and Na+ balance, which helped parasites survive in a relatively high concentration of CIP. An illustration of the putative processes for regulating Na+ and H+ in erythrocytes infected with the Babesia parasite is presented in Figure supplement 4.

In a previous study’s isolated membrane assay, CIP was found to be more effective at inhibiting parasite development than PfATP4-associated ATPase activity and reduced sensitivity of PfATP4-associated ATPase activity to the medication was associated with a decrease in parasite susceptibility to CIP (Rosling et al., 2018). In our study, the higher resistance to CIP-mediated inhibition of parasite proliferation observed in BgATP4L921I compared to BgATP4L921V correlated with a higher resistance of BgATP4-associated ATPase activity to the drug. These findings provide additional evidence that the ATPase activity measured in this study corresponds to the activity of the BgATP4 protein.

According to studies using the PfATP4 model for molecular docking, the G358S mutation results in a steric clash that lowers CIP’s binding affinity (Qiu et al., 2022). The results from the current study were similar to the PfATP4 model. The molecular docking was constructed using a ColabFold model of wild-type BgATP4, which predicted the binding mode and affinity between ATP4 protein and the ligand to provide a possible mechanistic explanation. It suggested that the L921V mutation caused changes at the atomic level, whereas the L921I mutation created a steric clash that reduced the binding affinity of CIP. The predicted affinity score of the L921I mutation was lower than that of the L921V mutation. Thus, it is possible that CIP had a weaker binding to BgATP4L921I than to BgATP4L921V. These findings are consistent with the results obtained from measuring the IC50 of the drug against parasites with the L921V and L921I mutations.

Earlier version: While CIP shows promise as an antibabesial agent, the emergence of resistance emphasizes the importance of rational drug use and the development of combination therapies. Our results demonstrate that CIP did not exhibit cross-resistance with ATO or TQ, as IC50 values for ATO and TQ in resistant strains showed minimal changes (<1-fold) compared to the wild-type strain. This observation suggests that CIP could be effectively combined with other anti-babesial drugs to reduce the risk of resistance development. Previous studies have also reported that a single dose of TQ may not be sufficient to prevent parasite relapse in immunocompromised hosts, recommending either repeated TQ administration or combination therapy with other antibabesial agents to address this issue (Liu et al., 2024). Supporting this approach, our study demonstrates that the combination of CIP plus TQ effectively cleared B. microti infections in SCID mice, with no detectable parasitemia or DNA observed at 90 DPI, indicating complete parasite elimination.

Revised version: While CIP shows promise as an antibabesial agent, the emergence of resistance emphasizes the importance of rational drug use and the development of combination therapies. Our results demonstrate that CIP did not exhibit cross-resistance with ATO or TQ, as IC50 values for ATO and TQ in resistant strains showed minimal changes (<1-fold) compared to the wild-type strain. This observation suggests that CIP could be effectively combined with other antibabesial drugs to reduce the risk of resistance development. Previous studies have also reported that a single dose of TQ may not be sufficient to prevent parasite relapse in immunocompromised hosts, recommending either repeated TQ administration or combination therapy with other antibabesial agents to address this issue (Mordue et al., 2019). Supporting this approach, our study demonstrates that the combination of CIP plus TQ effectively cleared B. microti infections in SCID mice, with no detectable parasitemia or DNA observed at 90 DPI, indicating complete parasite elimination.

This preclinical finding underscores the potential of the CIP plus TQ combination as a viable treatment strategy for babesiosis, especially in cases where monotherapy is ineffective. Further studies are needed to fully explore the pharmacokinetics, long-term efficacy, and potential side effects of such combinations in clinical settings, ultimately facilitating optimized therapeutic strategies.

In summary, our findings provide a comprehensive understanding of CIP’s efficacy, mechanism of resistance, and identifying strategies to enhance its therapeutic potential. The combination of CIP with drugs like TQ offers a promising approach to combat babesiosis and mitigate the development of resistance.

Materials and methods

Parasite culture

The parasites B. gibsoni Oita strain and B. bovis Texas strain were in vitro cultured in 24-well plates and maintained in an atmosphere of 5% CO2 and 5% O2 at 37 °C (Liu et al., 2018). For the in vivo studies, B. microti Peabody mjr strain-(ATCC PRA-99TM) and B. rodhaini Australia strain-infected RBCs (iRBCs), which were collected and diluted with phosphate-buffered saline (PBS) when the parasitemia levels in the donor mice reached ∼20% and 50%, respectively, and were intraperitoneally injected into BALB/c mice. Each BALB/c mouse was infected with 1.0 × 107 B. microti or B. rodhaini iRBCs for the in vivo trials (Ji et al., 2022a).

In vitro cytotoxicity of CIP and hemolysis rate in canine erythrocytes

Cultures of Madin-Darby canine kidney (MDCK) cells and human foreskin fibroblasts (HFF) were maintained at 37 °C under an atmosphere of 5% CO2 and 5% O2 and the cytotoxic effect of CIP (MedChem Express, Tokyo, Japan) was assessed using a cell viability assay by CCK-8 (Dojindo, Japan) as described previously (Li et al., 2023). The selectivity index is calculated as the ratio between the half maximal inhibitory concentration (IC50) and the 50% cytotoxic concentration (CC50) values.

Canine erythrocytes were collected from healthy beagle dogs raised in NRCPD and stocked in Vega y Martinez (VYM) phosphate-buffered saline solution at 4 °C (Vega et al., 1985). A canine erythrocyte hemolysis assay was performed at concentrations of 0.1, 1, 5, 10, 25, 50, and 100 µM as previously described (Ariefta et al., 2022).

Evaluation of the efficacy of CIP against Babesia parasites in vitro

The efficacy of CIP against B. gibsoni and B. bovis was determined using a fluorescence assay, as previously described (Guswanto et al., 2014). The IC50 values were determined from the fluorescence values and by non-linear regression analysis (curve fit) in GraphPad Prism 9 (GraphPad Software Inc., USA).

Chemotherapeutic effects of CIP against Babesia infections in vivo

CIP was evaluated on B. microti- and B. rodhaini-infected mice, as previously described (Ndayisaba et al., 2021; Tuvshintulga et al., 2022). When B. microti- and B. rodhaini-infected mice had a 1% average parasitemia at 4 and 2 days postinfection (DPI), respectively, the drug treatments were administered and continued for seven days. Three groups of B. microti-infected mice were administered different treatments. The CIP group (n = 6) and the ATO plus AZI group (n = 6) were orally treated with 20 mg/kg CIP and 20 mg/kg ATO plus 20 mg/kg AZI (Sigma, Tokyo, Japan), respectively. All drugs were prepared in sesame oil. The infected mice of the vehicle group (n = 6) orally received 0.2 mL of sesame oil as the control. Eighteen mice infected with B. rodhaini were likewise placed into three groups for treatment, and they received the same treatments as the mice infected with B. microti. A light microscope (Nikon, Japan) and a hematology analyzer (Celltac α MEK-6450, Nihon Kohden Corporation, Tokyo, Japan) were used to assess the parasitemia and hematocrit levels every two and four days, respectively.

Selection of CIP-resistant B. gibsoni in vitro

Selections were initiated by exposing six independent flasks, each containing 10 μL (5×106) B. gibsoni iRBCs mixed with 40 μL (4×108) RBCs into a 450 μL culture medium, which contained increasing concentrations of CIP: 5, 10, 20, 30 to 694 nM (10×IC50). The medium containing CIP was replaced daily until parasites treated with 10×IC50 CIP reached multiplication rates that were approximately comparable to those of the untreated controls (Hwang et al., 2010). Then, to evaluate the decreased sensitivity to CIP, IC50 values of resistant strains were determined by nonlinear regression using the GraphPad Prism software.

Detection of B. gibsoni ATP4 gene mutations

The genomic DNA of the mutant parasites was extracted and sequenced (Ji et al., 2022a). The primer sets used for sequencing are listed in Table supplement 1. The single nucleotide variants were identified by pairwise alignment to the BgATP4WT sequence (GenBank: JAVEPI010000002.1). Next-generation sequencing (NGS) was performed to detect the ratio of wild-type to mutant parasites by using the Illumina NovaSeq6000 sequencing platform (Jeon et al., 2021).

Morphological changes in CIP-treated in vitro cultured B. gibsoni

A microscopy assay was used to detect the morphological changes of wild-type B. gibsoni after exposure to 50 nM CIP for three consecutive days (Liu et al., 2021; Tayebwa et al., 2018). At 4 days post-treatment, ImageJ software was used to measure the sizes of 100 randomly selected parasites in the CIP-treated group and the control group on Giemsa-stained blood smears. After treating the parasites as described above, iRBCs were fixed in GA fixation buffer (2% glutaraldehyde in 0.1 M sodium cacodylate buffer containing 1 mM CaCl2 and 1 mM MgCl2), then stored in rinse buffer (0.1 M sodium cacodylate buffer containing 1 mM CaCl2 and 1 mM MgCl2) at 4 °C for transmission electron microscopy (TEM) analysis (Hidayati et al., 2023).

Parasites [Na+]i and pHi measurements

Earlier version: For both [Na+]i and pHi measurements, the wild-type and mutant parasites were initially separated from erythrocytes by treatment with saponin (0.05% [wt/vol]) for 5 min (Saliba and Kirk, 1999). The Na+-sensitive fluorescent dye SBFI (Thermo Fisher Scientific; product S1263) was used to quantify intracellular sodium ([Na+]i). Saponin-isolated parasites (at 1×108 parasites/mL) were loaded with SBFI (5 μM; in the presence of 0.02% w/v Pluronic F127) for 1 hr at room temperature (RT) (Spillman et al., 2013). Thereafter, SBFI-loaded parasites were resuspended in physiological saline (120 mM NaCl, 5 mM KCl, 25 mM HEPES, 20 mM D-glucose, and 1 mM MgCl2 [pH 7.1]) at RT in the presence or absence of CIP. A 96-well microtiter plate was filled with around 200 µL of parasite suspension per well. The dye-loaded cells fluoresced at 515 nm after being stimulated at 340 and 380 nm. Parasites loaded with SBFI were suspended in calibration buffers containing [Na+] values ranging from 0 to 140 mM (pH 7.1) to establish calibration curves (Diarra et al., 2001).

Revised version: For both [Na+]i and pHi measurements, the wild-type and mutant parasites were initially separated from erythrocytes by treatment with saponin (0.05% [wt/vol]) for 5 min (Saliba and Kirk, 1999). The Na+-sensitive fluorescent dye SBFI (Thermo Fisher Scientific; product S1263) was used to quantify intracellular sodium ([Na+]i). Saponin-isolated parasites (at 1×108 parasites/mL) were loaded with SBFI (5 μM; in the presence of 0.02% w/v Pluronic F127) for 1 hr at room temperature (RT) (Spillman et al., 2013). Thereafter, SBFI-loaded parasites were resuspended in physiological saline (120 mM NaCl, 5 mM KCl, 25 mM HEPES, 20 mM D-glucose, and 1 mM MgCl2 [pH 7.1]) at RT in the presence or absence of CIP. A 96-well microtiter plate was filled with around 200 µL of parasite suspension per well.

Added in the revised version: The dye-loaded cells fluoresced at 515 nm after being stimulated at 340 and 380 nm. Parasites loaded with SBFI were suspended in calibration buffers containing [Na+] values ranging from 0 to 140 mM (pH 7.1) to establish calibration curves (Diarra et al., 2001).

The cytosolic pH of wild-type and mutant strains was measured using the pH-sensitive fluorescent dye BCECF [2’,7’-bis-(2-carboxyethyl)-5-(and-6)- carboxyfluorescein] (Biotium; product 51011). The BCECF was added to saponin-isolated parasites by suspension (1×108 parasites/mL) and incubated for 20 minutes at 37 °C in RPMI-1640 culture medium (Gibco, USA) (Saliba and Kirk, 1999). Thereafter, the parasites loaded with dye were rinsed thrice (12,000 × g, 1 min) in RPMI-1640 culture medium, then resuspended in physiological saline (as previously mentioned) with or without various concentrations of CIP. A 96-well microtiter plate was filled with around 200 µL of parasite suspension per well. The dye-loaded cells fluoresced at 520 nm after being stimulated at 440 and 490 nm. A pH calibration was carried out for each experiment (Mohring et al., 2022).

Measurements of membrane ATPase activity

Membranes from isolated B. gibsoni were prepared by lysing the parasites in ice-cold deionized water containing 7 × Protease Inhibitor Cocktail Tablets (Roche, Germany). The membrane preparation was then washed three times with ice-cold deionized water, with protease inhibitors included in the first two washes (Qiu et al., 2022). Protein concentrations in the membrane samples were determined using a Bradford assay. The production of inorganic phosphate (Pi) from ATP hydrolysis was measured using the Malachite Green Phosphate Assay (BioAssay Systems, USA). Membrane preparations were diluted in either a high Na+ solution (final reaction conditions: 150 mM NaCl, 20 mM KCl, 2 mM MgCl2, 50 mM Tris, pH 7.2) or a Na+-free solution (final reaction conditions: 150 mM choline chloride, 20 mM KCl, 2 mM MgCl2, 50 mM Tris, pH 7.2) to achieve a final protein concentration of 40 μg/mL. CIP was added at the concentrations specified in the respective figure legends. Reactions were conducted according to the manufacturer’s protocol provided with the assay kit.

Multiple ATP4 sequence alignment and molecular docking

B. gibsoni ATP4 (GenBank: KAK1443404.1), B. bovis ATP4 (PiroplasmaDB: BBOV_IV010020), B. microti ATP4 (GenBank: BMR1_03g01005), T. gondii ATP4 (GenBank: XP_018635122.1), and Homo sapiens ATP4 (GenBank: NM_000704.3) sequences were obtained by a homology search using P. falciparum ATP4 (GenBank: PF3D7_1211900). Sequence alignment was analyzed using MUSCLE in Jalview v2.11.3.2 software and BLAST (http://www.ncbi.nlm.nih.gov/BLAST/).

AlphaFold was used to predict the structure of BgATP4WT (Jumper et al., 2021). Using the GROMACS 2021 Molecular Dynamics package, the energy minimization was carried out following the model’s generation (Lindahl et al., 2021). The mutations were produced by using Charmm-GUI PDB reader (Jo et al., 2014). PyMOL (version 2.0 Schrödinger, LLC) was used to confirm the position of the mutation site (center: -20.367, 10.904, 9.435; size: 30×30×30) (Trott and Olson, 2010). The ligand molecules CIP was downloaded from PubChem (CID 44469321 (https://pubchem.ncbi.nlm.nih.gov/compound/44469321). The ligand was used to dock with Gnina (Eberhardt et al., 2021). The affinity score and binding pose were chosen only from the highest CNN (convolutional neural network) score results from docking simulations. The models were visualized using the PyMOL Molecular Graphic System and Discovery Studio.

Cross-resistance to TQ and ATO in mutant parasites

The efficacy of TQ and ATO against BgATP4WT, BgATP4L921V, and BgATP4L921I was evaluated using a fluorescence assay, as described above. The IC50 values were calculated from the fluorescence values using non-linear regression analysis (curve fitting) in GraphPad Prism 9 (GraphPad Software Inc., USA).

Efficacy of CIP plus TQ combination therapy in SCID mice infected with B. microti

Twenty 6-week-old female immunocompromised mice (C.B-17/IcrJcl-Prkdcscid strain, CLEA Japan, Inc.) were intraperitoneally injected with 1.0 × 107 of B. microti-infected blood cells (Tuvshintulga et al., 2022). When the average parasitemia across all mice reached 1% (at 4 DPI), the drug treatments were initiated. The CIP group (n = 5) was orally administered 20 mg/kg CIP for seven days. The TQ group (n = 5) was orally administered a single dose of 20 mg/kg TQ. The CIP plus TQ group (n = 5) received the combination treatment, following the dosage and administration methods described above. The vehicle group (n = 5) orally received 0.2 mL of sesame oil as the control. Thin blood smears were prepared every other day and examined under a light microscope to determine parasitemia levels.

qPCR analysis

Earlier version: To further verify the presence or absence of B. microti DNA in treated SCID mice, real-time quantitative PCR analysis was performed as described previously (Vydyam et al., 2024). Briefly, genomic DNA was extracted from samples collected at 90 DPI and subjected to qPCR analysis using SYBR Green I, targeting a highly conserved region of Babesia mitochondrial genome (mtDNA). Primers used were: Bmic-F 5′-TTGCGATAGTAATAGATTTACTGC-3′ and B-lsu-R2 5′-TCTTAACCCAACTCACGTACCA-3′ (Qurollo et al., 2017). The reaction mixture consisted of: 1 × Advanced Universal SYBR Green Super Mix (2 ×) (1725270; Bio-Rad); 0.5 μM of each primer, and 2 μl of genomic DNA.

Revised version: To further verify the presence or absence of B. microti DNA in treated SCID mice, real-time quantitative PCR analysis was performed as described previously (Vydyam et al., 2024). Briefly, genomic DNA was extracted from blood samples (200 μl) collected at 90 DPI and subjected to qPCR analysis using SYBR Green I, targeting a highly conserved region of Babesia mitochondrial genome (mtDNA). Primers used were: Bmic-F 5′-TTGCGATAGTAATAGATTTACTGC-3′ and B-lsu-R2 5′-TCTTAACCCAACTCACGTACCA-3′ (Qurollo et al., 2017). The reaction mixture consisted of: 1 × Advanced Universal SYBR Green Super Mix (2 ×) (1725270; Bio-Rad); 0.5 μM of each primer, and 100 ng of genomic DNA.

Detection of B. microti ATP4 gene mutations in relapsed infected SCID mice

The genomic DNA of the mutant parasites was extracted and sequenced (Ji et al., 2022a). The primer sets used for sequencing are listed in Table supplement 2. The single nucleotide variants were confirmed by pairwise alignment to the BmATP4WT sequence (GenBank: BMR1_03g01005).

Statistical analysis

Data analysis, namely one-way analysis of variance (ANOVA) and two-tailed unpaired t-tests, was performed using GraphPad Prism (La Jolla, CA, USA) version 9. A P value of <0.05 was considered a statistically significant result.